---
  title: "مقدمة في الأساسيات"
  description: ""
---

## كيف تعمل

```yaml
type: NormalExercise 
lang: r
xp: 100 
skills: 1
key: 22b9e6bd8a   
```


في محرر على اليمين يجب عليك كتابة رمز R لحل التمارين. عندما تضغط على الزر "إرسال الإجابة" ، يتم تفسير كل سطر من الرموز وتنفيذها بواسطة R وستحصل على رسالة ما إذا كانت شفرتك صحيحة أم لا. يظهر إخراج رمز R في وحدة التحكم في الزاوية اليمنى السفلى.

يستفيد R من علامة # لإضافة تعليقات ، حتى تتمكن أنت والآخرين من فهم ما هو رمز R. تماما مثل تويتر! لا يتم تشغيل التعليقات كرمز R ، لذا لن تؤثر على نتائجك. على سبيل المثال ، حساب 3 + 4 في المحرر على اليمين هو تعليق.

يمكنك أيضا تنفيذ أوامر R مباشرة في وحدة التحكم. هذه طريقة جيدة لتجربة رمز R ، حيث لم يتم التحقق من إرسالك للتأكد من صحته.


`@instructions`
- في المحرر على اليمين هناك بالفعل بعض التعليمات البرمجية عينة. هل يمكنك معرفة أي من الأسطر رمز R الفعلي وأي تعليقات؟
- أضف سطرًا من الشفرة يحسب مجموع 6 و 12 ، ثم اضغط على الزر "إرسال رد".

`@hint`
- ما عليك سوى إضافة سطر من R code الذي يحسب مجموع 6 و 12 ، تمامًا مثل المثال في نموذج التعليمة البرمجية!

`@pre_exercise_code`

```{r}
# Load datasets and packages here.
```

`@sample_code`

```{r}
# احسب 3 + 4
3 + 4

# احسب 6 + 12

```

`@solution`

```{r}
# احسب 3 + 4
3 + 4

# احسب 6 + 12
6 + 12
```

`@sct`

```{r}
# Update this to something more informative.
success_msg("Some praise! Then reinforce a learning objective from the exercise.")
```
